### Codio Activity 16.9: Investigating your own data

For this activity, you are asked to go out and choose a dataset to build a classification model with.  Specifically, you are to compare the `LogisticRegression`, `KNearestNeighborsClassifier`, and `SVC` estimators in terms of performance and speed in model fitting.  You should optimize this model according to what metric you believe is the appropriate one for the task between `precision`, `recall`, or `accuracy`.  

```python
import pandas as pd
```

### Gathering the data

For your dataset, consider using an example dataset from either [kaggle](https://www.kaggle.com/) or the [UCI Machine Learning Repository](https://archive.ics.uci.edu/ml/index.php).  Select an appropriate dataset that is a classification problem.  Download the data file and work in a notebook locally to perform your analysis.  Be sure to grid search different model parameters and compare the different estimators.  Construct a DataFrame of the model results with the following information:

| model | train score | test score | average fit time |
| ----- | -----   | -------   | ------- |
| KNN | ? | ? | ? |
| Logistic Regression | ? | ? | ? |
| SVC | ? | ? | ? |

The assignment will expect a DataFrame with this exact structure and index and column names.  You will be graded based on the exact match of the structure of the DataFrame.  One suggestion is to build a DataFrame and write this out to `.json`, copy and paste this below to create the DataFrame.  Alternatively, you can write it out to a `.csv` file and copy the text, or simply hardcode the DataFrame based on your results.

### Problem 1

#### DataFrame of modeling results

Assign your constructed results DataFrame to `results_df` below.  Be sure that the `model` column above is the index of the DataFrame, and the three column names match the order and formatting of the example above.

```python
import pandas as pd
```

```python
df = pd.read_csv('Credit.csv')
df.head()
```

Output:
```
Unnamed: 0   Income  Limit  Rating  Cards  Age  Education  Gender Student  \
0           1   14.891   3606     283      2   34         11    Male      No   
1           2  106.025   6645     483      3   82         15  Female     Yes   
2           3  104.593   7075     514      4   71         11    Male      No   
3           4  148.924   9504     681      3   36         11  Female      No   
4           5   55.882   4897     357      2   68         16    Male      No   

  Married  Ethnicity  Balance  
0     Yes  Caucasian      333  
1     Yes      Asian      903  
2      No      Asian      580  
3      No      Asian      964  
4     Yes  Caucasian      331
```

<!-- grade_id: cell-c2a5ca74d8cdd115, points: 0, grade: false, solution: true -->
```python
### GRADED
results_df = ''

    
### BEGIN SOLUTION
res_dict = {'model': ['KNN', 'Logistic Regression', 'SVC'],
           'train score': [0, 0, 0],
           'test score': [0, 0, 0],
           'average fit time': [0, 0, 0]}
results_df = pd.DataFrame(res_dict).set_index('model')
### END SOLUTION

### ANSWER CHECK
print(type(results_df))
print(results_df.shape)
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
(3, 3)
```

<!-- grade_id: cell-a55e4dfa02723bd0, points: 50, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
res_dict_ = {'model': ['KNN', 'Logistic Regression', 'SVC'],
           'train score': [0, 0, 0],
           'test score': [0, 0, 0],
           'average fit time': [0, 0, 0]}
results_df_ = pd.DataFrame(res_dict_).set_index('model')
#
#
#
assert results_df.shape == results_df_.shape
assert type(results_df) == type(results_df_)
### END HIDDEN TESTS
```